<a href="https://colab.research.google.com/github/hema-demo1/Hemavathi/blob/main/Week06/week06_assignment_2.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Week 6 Exercises

_McKinney 6.1_

There are multiple ways to solve the problems below.  You can use any one of several approaches.  For example, you can read CSV files using Pandas or the csv module.  Your score won't depend on which modules you choose to use unless explicitly noted below, but your programming style will still matter.

### 30.0 Create sample data in local directory

```python
import pickle

raw = b'\x80\x04\x95\xac\x10\x00\x00\x00\x00\x00\x00X\xa5\x10\x00\x00{\n  "entry": [\n    {\n      "resource": {\n        "resourceType": "AllergyIntolerance",\n        "recordedDate": "2015-08-24T23:11:36Z",\n        "status": "confirmed",\n        "criticality": "CRITL",\n        "onset": "2012-11-07T00:00:00Z",\n        "patient": {\n          "display": "Jason Argonaut"\n        },\n        "substance": {\n          "text": "PENICILLIN G",\n          "coding": [\n            {\n              "system": "http://www.nlm.nih.gov/research/umls/rxnorm",\n              "code": "7980",\n              "display": "PENICILLIN G"\n            },\n            {\n              "system": "http://fdasis.nlm.nih.gov",\n              "code": "Q42T66VG0C",\n              "display": "PENICILLIN G"\n            }\n          ]\n        },\n        "reaction": [\n          {\n            "certainty": "confirmed",\n            "onset": "2012-11-07T00:00:00Z",\n            "manifestation": [\n              {\n                "text": "Hives"\n              }\n            ],\n            "note": {\n              "text": "Severity low enough to be prescribed if needed."\n            }\n          }\n        ],\n        "note": {\n          "text": "Severity low enough to be prescribed if needed."\n        }\n      }\n    },\n    {\n      "resource": {\n        "resourceType": "AllergyIntolerance",\n        "recordedDate": "2015-11-07T20:55:10Z",\n        "status": "confirmed",\n        "criticality": "CRITL",\n        "id": "TPcWiBG2h2E114Vh0sRT8fQB",\n        "onset": "2010-05-02T00:00:00Z",\n        "patient": {\n          "display": "Jason Argonaut"\n        },\n        "substance": {\n          "text": "SHELLFISH-DERIVED PRODUCTS",\n          "coding": [\n            {\n              "system": "http://hl7.org/fhir/ndfrt",\n              "code": "N0000007624",\n              "display": "SHELLFISH-DERIVED PRODUCTS"\n            }\n          ]\n        },\n        "reaction": [\n          {\n            "certainty": "confirmed",\n            "onset": "2010-05-02T00:00:00Z",\n            "manifestation": [\n              {\n                "text": "Itching"\n              }\n            ]\n          }\n        ]\n      }\n    },\n    {\n      "resource": {\n        "resourceType": "AllergyIntolerance",\n        "recordedDate": "2015-11-07T20:56:34Z",\n        "status": "confirmed",\n        "criticality": "CRITH",\n        "id": "TKebKfLXzu6Sp.LY-IpvpmQB",\n        "onset": "2014-03-07T00:00:00Z",\n        "patient": {\n          "display": "Jason Argonaut"\n        },\n        "substance": {\n          "text": "STRAWBERRY",\n          "coding": [\n            {\n              "system": "http://www.nlm.nih.gov/research/umls/rxnorm",\n              "code": "892484",\n              "display": "STRAWBERRY"\n            },\n            {\n              "system": "http://fdasis.nlm.nih.gov",\n              "code": "4J2TY8Y81V",\n              "display": "STRAWBERRY"\n            }\n          ]\n        },\n        "reaction": [\n          {\n            "certainty": "confirmed",\n            "onset": "2014-03-07T00:00:00Z",\n            "manifestation": [\n              {\n                "text": "Anaphylaxis"\n              }\n            ]\n          }\n        ]\n      }\n    },\n    {\n      "resource": {\n        "resourceType": "AllergyIntolerance",\n        "recordedDate": "2015-08-24T23:11:36Z",\n        "status": "confirmed",\n        "criticality": "CRITL",\n        "onset": "1978-12-08T00:00:00Z",\n        "patient": {\n          "display": "Paul Boal"\n        },\n        "substance": {\n          "text": "PENICILLIN G",\n          "coding": [\n            {\n              "system": "http://www.nlm.nih.gov/research/umls/rxnorm",\n              "code": "7980",\n              "display": "PENICILLIN G"\n            },\n            {\n              "system": "http://fdasis.nlm.nih.gov",\n              "code": "Q42T66VG0C",\n              "display": "PENICILLIN G"\n            }\n          ]\n        },\n        "reaction": [\n          {\n            "certainty": "confirmed",\n            "onset": "1978-12-08T00:00:00Z",\n            "manifestation": [\n              {\n                "text": "Bruising"\n              }\n            ],\n            "note": {\n              "text": "Do not prescribe!"\n            }\n          }\n        ],\n        "note": {\n          "text": "Do not prescribe!"\n        }\n      }\n    }\n  ]\n}\n\x94.'
json_data = pickle.loads(raw)

with open("allergies.json", "w") as f:
    f.write(json_data)
```

### 30.1 List of Allergies

After running 30.0 to create your input data file, there is a file called `allergies.json` that contains a list of patient allergies. It is taken from sample data provided by the EHR vendor, Epic, here: https://open.epic.com/Clinical/Allergy

Take some time to look at the structure of the file.  You can open it directly in Jupyter by clicking the _Home_ icon, then the _from_instructor_ folder, and then the _data_ folder.

Within the file, you'll see that it is a dictionary with many items in it.  One of those items is called `entry` and that item is a list of things.  You can tell that because the item name is immediately followed by an opening square bracket, signifying the start of a list.  It's line 11 of the file: `  "entry": [`

Write a function named `allergy_count(json_file)` that takes as one parameter the name of the JSON file and returns an integer number of entries in that file.  Your function should open the file, read the json into a Python object, and return how many items there are in the list of `entry`s.

```python
import json
ALLERGIES_FILE="allergies.json"
```

<!-- grade_id: q1-answer, points: 0, grade: false, solution: true -->
```python
# Put your solution here
def allergy_count(json_file):

    """
    This function reads and returns the total entries in the allergy data stored in JSON file.
     It scans the particular JSON file and reads the 'entry’ items and gives the number of entries.
     If the file can not be found or opened it will print an error message and then return 0.
     This function accepts one parameter which is json_file – should be name of the JSON file or its path.
     It returns an integer which is count of entries in the ‘entry’ list of the file.

    """
    try:
        with open(json_file, 'r') as allergy_data:
            json_content = json.load(allergy_data)
            if 'entry' in json_content:
                return len(json_content['entry'])
            else:
                return 0
    except FileNotFoundError:
        print(f"File '{json_file}' not found.")
        return

```

```python
allergy_count(ALLERGIES_FILE)
```

Output:
```
4
```

<!-- grade_id: q1-tests, points: 5, grade: true, solution: false -->
```python
assert type(allergy_count(ALLERGIES_FILE)) == int
assert allergy_count(ALLERGIES_FILE) == 4
```

### 30.2 Number of Patients

If you dig a little bit deaper into this list of allergies, you'll see that each result has a patient associated with it.  Create a funcation called `patient_count(json_file)` that will count how many unique patients we have in this JSON structure.  

```python
# Put your solution here
def patient_count(json_file):
    """
    Compute the frequency of distinct patients based from the given allergy JSON file.

    This function reads a JSON file and parsed respective patient information from any record.
    It then finds the number of tables returned for a given number of different patients.
    If the specified JSON file is not found or cannot be opened, it will print an error message exit with 0.

    Args:
    - json_file (str): The name or the path of that JSON file which needs to be processed.

    Returns:
    - int: The number of patients we have found in the JSON that are not previously seen.

    Examples:
        to reflect the patient count, I used the following type of code: patient_count(“ALLERGIES_FILE”)
        int
        It’s worth calling patient_count(ALLERGIES_FILE)
        2
    """
    try:
        with open(json_file, 'r') as file:
            allergy_records = json.load(file)
            distinct_patients = set()
            for allergy_entry in allergy_records.get('entry', []):
                patient_display_name = allergy_entry.get('resource', {}).get('patient', {}).get('display', '')
                if patient_display_name:
                    distinct_patients.add(patient_display_name)
            return len(distinct_patients)
    except FileNotFoundError:
        print(f"File '{json_file}' not found.")
        return 0

```

```python
patient_count(ALLERGIES_FILE)
```

Output:
```
2
```

```python
assert type(patient_count(ALLERGIES_FILE)) == int
assert patient_count(ALLERGIES_FILE) == 2
```

### 30.3 How Many Allergies per Patient

Although each entry is a separate allergy, several of them are for the same patient.  Write a function called `allergy_per_patient(json_file)` that counts up how many allergies each patient has.


```python
# Put your solution here
def allergy_per_patient(json_file):
    """
    Calculate the number of allergies per patient from a JSON file containing allergy data.
    This function reads a JSON file, extracts patient and allergy information from each record,
    and creates a dictionary where keys represent patient names and values represent the count of allergies for each patient.
    The function takes one argument, json_file, which is the name or path of the JSON file to be processed.
    It returns a dictionary where keys are patient names and values are the allergy counts for each patient.
    If the specified JSON file is not found or cannot be opened, it will display an error message and return an empty dictionary.
    Examples include using the function with a JSON file path to return the type as dict and a sample output showing the allergy counts.
    """
    try:
        with open(json_file, 'r') as file:
            data = json.load(file)

            allergies_per_patient = {}

            for record in data.get('entry', []):
                patient_name = record.get('resource', {}).get('patient', {}).get('display', '')

                if patient_name:
                    allergies_per_patient[patient_name] = allergies_per_patient.get(patient_name, 0) + 1

            return allergies_per_patient
    except FileNotFoundError:
        print(f"File '{json_file}' not found.")
        return {}

```

```python
allergy_per_patient(ALLERGIES_FILE)
```

Output:
```
{'Jason Argonaut': 3, 'Paul Boal': 1}
```

```python
assert type(allergy_per_patient(ALLERGIES_FILE)) == dict
assert allergy_per_patient(ALLERGIES_FILE) == {'Paul Boal': 1, 'Jason Argonaut': 3}
```

### 30.4 Patient Allergies and Reaction

You'll see in the file that each of the items in the `entry` list have several other attributes including a patient name, substance text representation, and a reaction manifestation.  Create a function named `allergy_list(json_file)` that will create an output list that has patient name, allergy, and reaction for each `entry`.  The actual result you should get will be:

```python
[['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis']]
```

You'll notice that the reaction and the manifestation of that action are lists.  You only need to capture the first reaction and the first manifestation of the action.  That is, if there is a list of things, just output the first one.

<!-- grade_id: q2-answer, points: 0, grade: false, solution: true -->
```python
# Put your solution here
def allergy_list(json_file):
    """

    This function extracts allergy data from a JSON file that contains information about patient allergies.
    It reads the specified file and gathers details on patients, their allergies, and any associated reactions.
    The result is returned as a list of lists, where each inner list includes a patient's name, the allergy,
    and the reaction (if there is one). If the file can't be found or opened, the function will simply return an empty list.

    Args:
    - json_file (str): The path or name of the JSON file containing allergy information.

    Returns:
    - list: A collection of lists that provide patient names, their allergies, and any reactions.
    """
    try:
        with open(json_file, 'r') as file:
            allergy_data = json.load(file)
            results = []

            entries = allergy_data.get("entry", [])
            for entry in entries:
                resource_info = entry.get("resource", {})
                patient = resource_info.get("patient", {}).get("display", None)
                allergy_substance = resource_info.get("substance", {}).get("text", None)

                reaction_info = None
                reactions_list = resource_info.get("reaction", [])
                if reactions_list:
                    first_reaction_detail = reactions_list[0].get("manifestation", [])
                    if first_reaction_detail:
                        reaction_info = first_reaction_detail[0].get("text", None)

                results.append([patient, allergy_substance, reaction_info])

            return results
    except FileNotFoundError:
        return []

```

```python
allergy_list(ALLERGIES_FILE)
```

Output:
```
[['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising']]
```

<!-- grade_id: q2-tests, points: 10, grade: true, solution: false -->
```python
assert allergy_list(ALLERGIES_FILE) == [['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising']]

```

### 30.5 Allergy Reaction

Write a function called `allergy_reaction(json_file,patient,substance)` that takes three parameter and returns the reaction that will happen if the patient takes the specified substance.  You can solve this, in part, by calling your `allergy_list` function inside your new `allergy_reaction` function.

If the substance is not found in the allergy list, the function should return None.

<!-- grade_id: q3-answer, points: 0, grade: false, solution: true -->
```python
# Put your solution here
def allergy_reaction(json_file, patient, substance):
    """
   Finding out whether the patient is allergic to that substance or not.
    This function is to extract the allergy data from a file with JSON format and try to find a match
    between the given patient and the substances. If a match is found, then it returns
    the allergic reaction associated with it; otherwise, it returns None.

    Args:
    - json_file (str): The path for the JSON file containing the list of allergies.
    - patient (str): The ID number of the patient whose reaction is being tested.
    - substance (str): The substance to examine for an allergic reaction.

    Returns:
    - str or None: The allergic reaction if found; otherwise, Not applicable.

    Examples:
        allergy_reaction(“ALLERGIES_FILE.json”, “John Doe”, “Peanuts”)
        ‘If given Peanuts, they may have a reaction, such as itchy skin.’
        The function is calling the file name ‘ALLERGIES_FILE.json’ and two inputs: ‘Alice Smith’, Lactose

    """
    allergy_info = allergy_list(json_file)

    allergic_reaction = None
    for record in allergy_info:
        record_patient, record_substance, record_reaction = record
        if record_patient == patient and record_substance == substance:
            allergic_reaction = record_reaction
    return allergic_reaction

```

```python
allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN G')
```

Output:
```
'Hives'
```

<!-- grade_id: q3-tests, points: 5, grade: true, solution: false -->
```python
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN G') == 'Hives'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS') == 'Itching'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'STRAWBERRY') == 'Anaphylaxis'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN') == None
assert allergy_reaction(ALLERGIES_FILE, 'Paul Boal', 'PENICILLIN G') == 'Bruising'
```

---

## Check your work above

If you didn't get them all correct, take a few minutes to think through those that aren't correct.


## Submitting Your Work

In order to submit your work, you'll need to save this notebook file back to GitHub.  To do that in Google Colab:
1. File -> Save a Copy in GitHub
2. Make sure your HDS5210 repository is selected
3. Make sure the file name includes the week number like this: `week06/week06_assignment_2.ipynb`
4. Add a commit message that means something

**Be sure week names are lowercase and use a two digit week number!!**

**Be sure you use the same file name provided by the instructor!!**

